## Imports

```python
# pip install --upgrade scikit-learn
```

```python
import time
from IPython.display import clear_output
import numpy    as np
import pandas   as pd
import seaborn  as sb
import matplotlib.pyplot as plt
import sklearn  as skl

from sklearn import pipeline      # Pipeline
from sklearn import preprocessing # OrdinalEncoder, LabelEncoder
from sklearn import impute
from sklearn import compose
from sklearn import model_selection # train_test_split
from sklearn import metrics         # accuracy_score, balanced_accuracy_score, plot_confusion_matrix
from sklearn import set_config

set_config(display='diagram') # Useful for display the pipeline

print("Pandas  ", pd.__version__)
print("Sklearn ", skl.__version__) # Try to use 0.24
```

Output:
```
Pandas   1.3.5
Sklearn  0.24.2
```

## Get the dataset
- **CLOUD = True**: Download dataset from Kaggle. Necesary for cloud enviroments like COLAB. **Specify your [kaggle credentials](https://www.kaggle.com/docs/api)**.
- **CLOUD = False**: Get the dataset from your local machine. **Specify the data path**.

```python
DATA_PATH = "D:/Machine-Learning_DS/Chapter 02/09. Robust ML/titanic/"
```

## Load data

```python
df      = pd.read_csv(DATA_PATH + "train.csv", index_col='PassengerId')
df_test = pd.read_csv(DATA_PATH + "test.csv",  index_col='PassengerId')

print("Train DataFrame:", df.shape)
print("Test DataFrame: ", df_test.shape)
```

Output:
```
Train DataFrame: (891, 11)
Test DataFrame:  (418, 10)
```

```python
df
```

Output:
```
Survived  Pclass  \
PassengerId                     
1                   0       3   
2                   1       1   
3                   1       3   
4                   1       1   
5                   0       3   
...               ...     ...   
887                 0       2   
888                 1       1   
889                 0       3   
890                 1       1   
891                 0       3   

                                                          Name     Sex   Age  \
PassengerId                                                                    
1                                      Braund, Mr. Owen Harris    male  22.0   
2            Cumings, Mrs. John Bradley (Florence Briggs Th...  female  38.0   
3                                       Heikkinen, Miss. Laina  female  26.0   
4                 Futrelle, Mrs. Jacques Heath (Lily May Peel)  female  35.0   
5                                     Allen, Mr. William Henry    male  35.0   
...                    
```

## Check missings

```python
df.isnull().sum()
```

Output:
```
Survived      0
Pclass        0
Name          0
Sex           0
Age         177
SibSp         0
Parch         0
Ticket        0
Fare          0
Cabin       687
Embarked      2
dtype: int64
```

```python
df_test.isnull().sum()
```

Output:
```
Pclass        0
Name          0
Sex           0
Age          86
SibSp         0
Parch         0
Ticket        0
Fare          1
Cabin       327
Embarked      0
dtype: int64
```

```python
df
```

Output:
```
Survived  Pclass  \
PassengerId                     
1                   0       3   
2                   1       1   
3                   1       3   
4                   1       1   
5                   0       3   
...               ...     ...   
887                 0       2   
888                 1       1   
889                 0       3   
890                 1       1   
891                 0       3   

                                                          Name     Sex   Age  \
PassengerId                                                                    
1                                      Braund, Mr. Owen Harris    male  22.0   
2            Cumings, Mrs. John Bradley (Florence Briggs Th...  female  38.0   
3                                       Heikkinen, Miss. Laina  female  26.0   
4                 Futrelle, Mrs. Jacques Heath (Lily May Peel)  female  35.0   
5                                     Allen, Mr. William Henry    male  35.0   
...                    
```

# Exercise 1 (2pts):
Extract the title (Mr, Mrs, ... ) from the "Name" column.

Tips:
- split(',')[1] to get the 2nd part, and remove the surnamename
- split('.')[0] to get the 1str part, and remove the name

<!-- grade_id: cell-d56107842f0f54d2, points: 0, grade: false, solution: true -->
```python
# CODE HERE get_Title_from_Name funtion
# Create this function using lambda (not def)

get_Title_from_Name = lambda name: name.split(',')[1].split('.')[0].strip()


```

<!-- grade_id: cell-f3d989a17018344b, points: 2, grade: true, solution: false -->
```python
# assert df['Title'].values[0] == "Mr"
# assert df['Title'].values[1] == "Mrs"
# assert df['Title'].values[2] == "Miss"

# assert df_test['Title'].values[0] == "Mr"
# assert df_test['Title'].values[1] == "Mrs"
# assert df_test['Title'].values[414] == "Dona"
```

# Exercise 2 (1pts):
Apply the title_dictionary to get a better information about the title. You have to overwrite the Title variable.

```python
title_dictionary = {
    "Capt": "Officer",
    "Col": "Officer",
    "Major": "Officer",
    "Jonkheer": "Royalty",
    "Don": "Royalty",
    "Sir" : "Royalty",
    "Dr": "Officer",
    "Rev": "Officer",
    "the Countess":"Royalty",
    "Mme": "Mrs",
    "Mlle": "Miss",
    "Ms": "Mrs",
    "Mr" : "Mr",
    "Mrs" : "Mrs",
    "Miss" : "Miss",
    "Master" : "Master",
    "Lady" : "Royalty"
}
```

<!-- grade_id: cell-5c6a842c812fafda, points: 0, grade: false, solution: true -->
```python
# Use map to apply the prevous dict

df['Title']      = df['Name'].map(get_Title_from_Name)
df_test['Title'] = df_test['Name'].map(get_Title_from_Name)


```

<!-- grade_id: cell-669f5a637e57e835, points: 1, grade: true, solution: false -->
```python
# assert df['Title'].values[886] == "Officer"
# assert df_test['Title'].values[417] == "Master"
```

# Exercise OPTINAL (0pts):
Try to extract some information from the feature **Ticket**. Search on Internet if that colum has some kind of information.

```python
df['Ticket'].describe()
```

Output:
```
count        891
unique       681
top       347082
freq           7
Name: Ticket, dtype: object
```

# Exercise OPTIONAL (0pts):
Try to extract some information from the feature **Cabin**. Search on Internet if that colum has some kind of information.

```python
df.Cabin.describe()
```

Output:
```
count         204
unique        147
top       B96 B98
freq            4
Name: Cabin, dtype: object
```

# Preprocessing
For X data, notice that...
- We drop Survived because is the target variable
- We drop Name because we have extracted the Title: Mr, Mrs, ...
- We drop Ticket because it has no information -> see df.Ticket.nunique()
- We drop Cabin because it has a lot of missings (77% are missings)

Then, we identify **numerical** variables and **categorical** variables,

```python
x = df.drop(columns=["Survived", 'Name', 'Ticket', 'Cabin']) # X DATA (WILL BE TRAIN+VALID DATA)
y = df["Survived"] # 0 = No, 1 = Yes

x_test = df_test.drop(columns=['Name', 'Ticket', 'Cabin']) # # X_TEST DATA (NEW DATA)
```

```python
cat_vars  = ['Sex', 'Embarked', 'Title']         # x.select_dtypes(include=[object]).columns.values.tolist()
num_vars  = ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age'] # x.select_dtypes(exclude=[object]).columns.values.tolist()

print("\nNumerical features:\n", num_vars)
print("\nCategorical features:\n", cat_vars)
```

Output:
```
Numerical features:
 ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age']

Categorical features:
 ['Sex', 'Embarked', 'Title']
```

# Exercise 3 (2pts):
Create a **ColumnTransformer for Tree Models**. You need to create 2 pipelines (one for numerical and other for categories). Remember:
- Categorical pipeline: Some SimpleImputer -> Some Encoder
- Numerical pipeline: Some SimpleImputer -> NO Encoder

<!-- grade_id: cell-c607e75fb38ea248, points: 0, grade: false, solution: true -->
```python

num_preprocessing = pipeline.Pipeline(steps=[
  ('imputer', impute.SimpleImputer(strategy='mean'))
  # Some SimpleImputer here

])

cat_preporcessing = pipeline.Pipeline(steps=[
  ('imputer', impute.SimpleImputer(strategy='constant', fill_value='missing')),
  ('ordinal', preprocessing.OrdinalEncoder(handle_unknown='use_encoded_value', unknown_value = -1))
  # Some SimpleImputer here
  # Some Encoder here. Remember to handle_unknown
])

tree_prepro = compose.ColumnTransformer(transformers=[
    ('num', num_preprocessing, num_vars),
    ('cat', cat_preporcessing, cat_vars),
], remainder='drop') # Drop other vars not specified in num_vars or cat_vars

tree_prepro


### BEGIN SOLUTION

```

Output:
```
ColumnTransformer(transformers=[('num',
                                 Pipeline(steps=[('imputer', SimpleImputer())]),
                                 ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age']),
                                ('cat',
                                 Pipeline(steps=[('imputer',
                                                  SimpleImputer(fill_value='missing',
                                                                strategy='constant')),
                                                 ('ordinal',
                                                  OrdinalEncoder(handle_unknown='use_encoded_value',
                                                                 unknown_value=-1))]),
                                 ['Sex', 'Embarked', 'Title'])])
```

<!-- grade_id: cell-e636558135fb5975, points: 2, grade: true, solution: false -->
```python
assert type(tree_prepro)      is compose._column_transformer.ColumnTransformer
assert type(num_preprocessing) is pipeline.Pipeline
assert type(cat_preporcessing) is pipeline.Pipeline
assert len(num_preprocessing) == 1
assert len(cat_preporcessing) == 2
```

# Exercise 4 (1pts):
1. Complete the diccionary with some Tree Models.
2. Then we put each model in a Pipeline where:
   - first is the prepocessing with the column Transformer
   - Then is the Tree model
3. Display the fullpipeline of the LGBMClassifier

```python
from sklearn.tree          import DecisionTreeClassifier
from sklearn.ensemble      import RandomForestClassifier
from sklearn.ensemble      import ExtraTreesClassifier
from sklearn.ensemble      import AdaBoostClassifier
from sklearn.ensemble      import GradientBoostingClassifier
from sklearn.experimental  import enable_hist_gradient_boosting # Necesary for HistGradientBoostingClassifier
from sklearn.ensemble      import HistGradientBoostingClassifier
from xgboost               import XGBClassifier
from lightgbm              import LGBMClassifier
from catboost              import CatBoostClassifier
```

<!-- grade_id: cell-76a3be8223730c5a, points: 0, grade: false, solution: true -->
```python

tree_classifiers = {
  "Decision Tree": DecisionTreeClassifier(),
  "Extra Trees": ExtraTreesClassifier(),
  "Random Forest" : RandomForestClassifier(),
  "AdaBoost": AdaBoostClassifier(),
  "Skl GBM":  GradientBoostingClassifier(),
  "Skl HistGBM": HistGradientBoostingClassifier(),
  "XGBoost": XGBClassifier(),
  "LightGBM": LGBMClassifier(),
  "CatBoost": CatBoostClassifier()

  }
tree_classifiers = {name: pipeline.make_pipeline(tree_prepro, model) for name, model in tree_classifiers.items()}
tree_classifiers["LightGBM"]




```

Output:
```
Pipeline(steps=[('columntransformer',
                 ColumnTransformer(transformers=[('num',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer())]),
                                                  ['Pclass', 'SibSp', 'Parch',
                                                   'Fare', 'Age']),
                                                 ('cat',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer(fill_value='missing',
                                                                                 strategy='constant')),
                                                                  ('ordinal',
                                                                   OrdinalEncoder(handle_unknown='use_encoded_value',
                                                       
```

<!-- grade_id: cell-d4744022b37e2e9b, points: 1, grade: true, solution: false -->
```python
for pipe in tree_classifiers.values():
    assert type(pipe) is pipeline.Pipeline
```

# Exercise 5 (3pts):
Define a simple split validation strategy with:
- 80% for train
- 20% for validation
- With stratification
- random_state=0

And train all the models in a for loop

<!-- grade_id: cell-64f17e0d448bca7e, points: 0, grade: false, solution: true -->
```python

x_train, x_val, y_train, y_val = model_selection.train_test_split(
    x, y,  test_size= 0.2, stratify = y,  random_state= 0
)


results = pd.DataFrame({'Model': [], 'Accuracy': [], 'Bal Acc.': [], 'Time': []})

for model_name, model in tree_classifiers.items():
    start_time = time.time()
    model.fit(x_train, y_train)
    
    # FOR EVERY PIPELINE (PREPRO + MODEL) -> TRAIN WITH TRAIN DATA (x_train)
    
    # GET PREDICTIONS USING x_val
    pred = model.predict(x_val)

    total_time = time.time() - start_time

    results = results.append({"Model":    model_name,
                              "Accuracy": metrics.accuracy_score(y_val, pred)*100,
                              "Bal Acc.": metrics.balanced_accuracy_score(y_val, pred)*100,
                              "Time":     total_time},
                              ignore_index=True)
                              
                              


### BEGIN SOLUTION




results_ord = results.sort_values(by=['Accuracy'], ascending=False, ignore_index=True)
results_ord.index += 1 
results_ord.style.bar(subset=['Accuracy', 'Bal Acc.'], vmin=0, vmax=100, color='#5fba7d')
```

Output:
```
Learning rate set to 0.008911
0:	learn: 0.6896767	total: 176ms	remaining: 2m 55s
1:	learn: 0.6840123	total: 181ms	remaining: 1m 30s
2:	learn: 0.6786433	total: 186ms	remaining: 1m 1s
3:	learn: 0.6729133	total: 189ms	remaining: 47.1s
4:	learn: 0.6662457	total: 195ms	remaining: 38.8s
5:	learn: 0.6608927	total: 199ms	remaining: 33s
6:	learn: 0.6551087	total: 203ms	remaining: 28.8s
7:	learn: 0.6506842	total: 205ms	remaining: 25.5s
8:	learn: 0.6454167	total: 211ms	remaining: 23.3s
9:	learn: 0.6404293	total: 216ms	remaining: 21.3s
10:	learn: 0.6351048	total: 221ms	remaining: 19.8s
11:	learn: 0.6296989	total: 225ms	remaining: 18.6s
12:	learn: 0.6258990	total: 228ms	remaining: 17.3s
13:	learn: 0.6210795	total: 232ms	remaining: 16.3s
14:	learn: 0.6167380	total: 236ms	remaining: 15.5s
15:	learn: 0.6130757	total: 240ms	remaining: 14.8s
16:	learn: 0.6075493	total: 244ms	remaining: 14.1s
17:	learn: 0.6028486	total: 247ms	remaining: 13.5s
18:	learn: 0.5986144	total: 251ms	remaining: 12.9s
19:	learn: 
```

<!-- grade_id: cell-af27b33fe88ad384, points: 3, grade: true, solution: false -->
```python
assert results_ord["Accuracy"].min() > 75
assert results_ord["Bal Acc."].min() > 75
assert len(results_ord) == 9
```

Output:
```
AssertionError                            Traceback (most recent call last)
C:\Users\VARUNK~1\AppData\Local\Temp/ipykernel_2700/3542778932.py in <module>
      1 assert results_ord["Accuracy"].min() > 75
----> 2 assert results_ord["Bal Acc."].min() > 75
      3 assert len(results_ord) == 9

AssertionError:
```

# Exercise 6 (3pts):
Define a 10 Fold cross validation strategy with:
- With stratification
- shuffle=True
- random_state=0

And train all the models in a for loop.

Tip you can use **[cross_val_predict](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.cross_val_predict.html)** for both training and predict with 

<!-- grade_id: cell-37bf16dc25b43732, points: 0, grade: false, solution: true -->
```python

skf = model_selection.StratifiedKFold(
    n_splits= 10, shuffle=True, random_state=0
)




results = pd.DataFrame({'Model': [], 'Accuracy': [], 'Bal Acc.': [], 'Time': []})


for model_name, model in tree_classifiers.items():
    start_time = time.time()
   
        
    # TRAIN AND GET PREDICTIONS USING cross_val_predict() and x,y
    pred = model_selection.cross_val_predict(model, x, y, cv = skf)

    total_time = time.time() - start_time

    results = results.append({"Model":    model_name,
                              "Accuracy": metrics.accuracy_score(y, pred)*100,
                              "Bal Acc.": metrics.balanced_accuracy_score(y, pred)*100,
                              "Time":     total_time},
                              ignore_index=True)
                              
                              



results_ord = results.sort_values(by=['Accuracy'], ascending=False, ignore_index=True)
results_ord.index += 1 
results_ord.style.bar(subset=['Accuracy', 'Bal Acc.'], vmin=0, vmax=100, color='#5fba7d')
```

Output:
```
Learning rate set to 0.009371
0:	learn: 0.6897520	total: 11.3ms	remaining: 11.3s
1:	learn: 0.6846201	total: 12.8ms	remaining: 6.38s
2:	learn: 0.6795991	total: 14.7ms	remaining: 4.87s
3:	learn: 0.6738699	total: 17.2ms	remaining: 4.29s
4:	learn: 0.6683363	total: 18.8ms	remaining: 3.75s
5:	learn: 0.6626931	total: 20.7ms	remaining: 3.43s
6:	learn: 0.6568739	total: 22.8ms	remaining: 3.23s
7:	learn: 0.6522961	total: 24.6ms	remaining: 3.05s
8:	learn: 0.6465951	total: 27ms	remaining: 2.98s
9:	learn: 0.6410359	total: 29.3ms	remaining: 2.9s
10:	learn: 0.6358044	total: 31.3ms	remaining: 2.82s
11:	learn: 0.6303996	total: 33.4ms	remaining: 2.75s
12:	learn: 0.6249874	total: 35.7ms	remaining: 2.71s
13:	learn: 0.6206621	total: 37.2ms	remaining: 2.62s
14:	learn: 0.6168637	total: 38.3ms	remaining: 2.51s
15:	learn: 0.6125757	total: 40.3ms	remaining: 2.48s
16:	learn: 0.6086068	total: 42.8ms	remaining: 2.48s
17:	learn: 0.6043199	total: 45.4ms	remaining: 2.48s
18:	learn: 0.5998418	total: 47.4ms	remaining: 2
```

<!-- grade_id: cell-57cc085faad513cf, points: 3, grade: true, solution: false -->
```python
assert results_ord["Accuracy"].min() > 75
assert results_ord["Bal Acc."].min() > 75
assert len(results_ord) == 9
```

```python
results_ord.iloc[0][0]
```

Output:
```
'CatBoost'
```

# Exercise 7.1
Train with all data the best model

```python
# best_model = tree_classifiers["SELECT MY BEST MODEL HERE"]
best_model = tree_classifiers[results_ord.iloc[0][0]]
best_model

# Fit best model with all data


```

Output:
```
Pipeline(steps=[('columntransformer',
                 ColumnTransformer(transformers=[('num',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer())]),
                                                  ['Pclass', 'SibSp', 'Parch',
                                                   'Fare', 'Age']),
                                                 ('cat',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer(fill_value='missing',
                                                                                 strategy='constant')),
                                                                  ('ordinal',
                                                                   OrdinalEncoder(handle_unknown='use_encoded_value',
                                                       
```

```python
best_model.fit(x,y)
```

Output:
```
Learning rate set to 0.009807
0:	learn: 0.6861635	total: 16.7ms	remaining: 16.7s
1:	learn: 0.6802762	total: 20.9ms	remaining: 10.4s
2:	learn: 0.6745235	total: 24.3ms	remaining: 8.09s
3:	learn: 0.6693773	total: 26.4ms	remaining: 6.57s
4:	learn: 0.6628686	total: 29.6ms	remaining: 5.88s
5:	learn: 0.6572098	total: 32.4ms	remaining: 5.36s
6:	learn: 0.6516398	total: 36.3ms	remaining: 5.15s
7:	learn: 0.6455882	total: 38.8ms	remaining: 4.82s
8:	learn: 0.6396866	total: 41.5ms	remaining: 4.57s
9:	learn: 0.6342669	total: 45.7ms	remaining: 4.52s
10:	learn: 0.6282907	total: 48.2ms	remaining: 4.34s
11:	learn: 0.6224250	total: 50.7ms	remaining: 4.17s
12:	learn: 0.6178670	total: 52.5ms	remaining: 3.99s
13:	learn: 0.6138446	total: 55.3ms	remaining: 3.9s
14:	learn: 0.6112188	total: 56.8ms	remaining: 3.73s
15:	learn: 0.6072904	total: 59.9ms	remaining: 3.68s
16:	learn: 0.6035166	total: 62.2ms	remaining: 3.6s
17:	learn: 0.5989585	total: 64.8ms	remaining: 3.54s
18:	learn: 0.5948441	total: 67.4ms	remaining: 
```

# Exercise 7.2 (2pts)
With your best model, generate the predicitions for test data (x_test)

<!-- grade_id: cell-44b93a0dbd4eb6fc, points: 0, grade: false, solution: true -->
```python
# test_pred = # Get the predictions for x_test

test_pred = best_model.predict(x_test)


```

<!-- grade_id: cell-a29fc8d24284520e, points: 2, grade: true, solution: false -->
```python
assert len(test_pred) == 418
assert np.unique(test_pred).tolist() == [0,1]
```

# Exercise 7.3

Submit to kaggle.

- You can use the kaggle command line app. Check https://github.com/Kaggle/kaggle-api

```python
sub = pd.DataFrame(test_pred, index=x_test.index, columns=["Survived"])
sub.head()
```

Output:
```
Survived
PassengerId          
892                 0
893                 0
894                 0
895                 0
896                 0
```

```python
sub.to_csv("sub.csv")

```

```python
!kaggle competitions submit -c titanic -f sub.csv -m "My submission message"
```

Output:
```
'kaggle' is not recognized as an internal or external command,
operable program or batch file.
```